# CS211: Data Privacy
## Homework 7

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1 (5 points)

(Reference [Chapter 7](https://uvm-plaid.github.io/programming-dp/notebooks/ch7.html) of the textbook)

Consider the following mean query:

```python
def mean_age():
    clipped_ages = adult['Age'].clip(lower=0, upper=100)
    return clipped_ages.mean()

mean_age()
```

Implement `ls_mean_at_distance`, an upper bound on the local sensitivity of the `mean_age` query at distance $k$, and `dist_to_high_ls_mean`, an upper bound on the distance from the true dataset to one with local sensitivity greater than or equal to $s_p$.

<!-- grade_id: cell-63f79c9bbc970326, points: 5, grade: true, solution: true -->
```python
def ls_mean_at_distance(k):
    # YOUR CODE HERE
    raise NotImplementedError()

def dist_to_high_ls_mean(s_p):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## Question 2 (10 points)

Implement `ptr_mean`, which should use the propose-test-release framework to calculate a differentially private estimate of the mean age. If the test fails, return `None`.

<!-- grade_id: cell-2c7d8b212cf7534c, points: 0, grade: false, solution: true -->
```python
def ptr_mean(s_p, epsilon, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

# proposed sensitivity: 0.05
# epsilon, delta = (1.0, 10^-5)
ptr_mean(0.05, 1.0, 1e-5)
```

<!-- grade_id: cell-050c04785d3ca701, points: 10, grade: true, solution: false -->
```python
# TEST CASE
true_mean = mean_age()
trials = [ptr_mean(0.05, 0.1, 1e-5) for _ in range(20)]
errors = [pct_error(true_mean, t) for t in trials]
assert np.mean(errors) < 10

assert ptr_mean(0.0001, 0.1, 1e-5) == None
```

## Question 3 (20 points)

Consider the `median_age` function, which calculates the *median* age (this version truncates if the length of the dataset is even), and the `ls_median` function, an upper bound on the local sensitivity of the median query:

```python
## Cache the sorted ages, because we will use them a lot.
sorted_ages = adult['Age'].clip(lower=0, upper=100).sort_values()

def median_age():
    idx = int(len(adult)/2)
    return sorted_ages.iloc[idx]

print('Median age:', median_age())

def ls_median():
    idx = int(len(adult)/2)
    return max(sorted_ages.iloc[idx] - sorted_ages.iloc[idx-1],
               sorted_ages.iloc[idx+1] - sorted_ages.iloc[idx])

print('Local sensitivity of the median:', ls_median())
```

Note that the local sensitivity of the median is 0. Implement the functions `ls_median_at_distance`, which calculates the local sensitivity at distance $k$ of the median query above, and the corresponding `dist_to_high_ls_median`.

*Hint*: note that the ages are clipped. Think about the worst-case scenario of adding or removing $k$ rows.

<!-- grade_id: cell-16a1fbaa4a589fd6, points: 0, grade: false, solution: true -->
```python
def ls_median_at_distance(k):
    # YOUR CODE HERE
    raise NotImplementedError()

def dist_to_high_ls_median(s_p):
    # YOUR CODE HERE
    raise NotImplementedError()

ls_median_at_distance(10000)
```

<!-- grade_id: cell-c7616dbe4cdc3fde, points: 20, grade: true, solution: false -->
```python
assert ls_median_at_distance(500) == 1
assert ls_median_at_distance(5000) == 6
assert ls_median_at_distance(10000) == 14
assert ls_median_at_distance(15000) == 28
```

## Question 4 (10 points)

Use the propose-test-release framework, plus `dist_to_high_ls_median`, to answer the median query with differential privacy.

<!-- grade_id: cell-8931f6c4576d8c47, points: 0, grade: false, solution: true -->
```python
def ptr_median(s_p, epsilon, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

ptr_median(0.01, 1.0, 1e-5)
```

<!-- grade_id: cell-90cffe527d83d7c7, points: 10, grade: true, solution: false -->
```python
# TEST CASE
true_median = median_age()
trials = [ptr_median(0.05, 0.1, 1e-5) for _ in range(20)]
errors = [pct_error(true_median, t) for t in trials]
assert np.mean(errors) < 10
```

## Question 5 (10 points)

In 2-5 sentences, answer the following:

- At roughly what distance does the local sensitivity of the median become non-zero?
- For what proposed sensitivity does `ptr_median` fail the test (i.e. return `None`)?
- What does this mean for the amount of noise required to release the median with differential privacy?

- At distance 400
- None! You can propose a sensitivity bound of 0
- This means you can release the median with *no noise* (for certain values of $\epsilon$)

## Question 6 (20 points)

Use the sample-and-aggregate framework to release the mean of the "Capital Gain" column. Reference [Chapter 7](https://uvm-plaid.github.io/programming-dp/notebooks/ch7.html).

<!-- grade_id: cell-de5a77cbd2059908, points: 0, grade: false, solution: true -->
```python
def f(chunk):
    return chunk.mean()

def saa_avg_capgain(k, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

saa_avg_capgain(500, 1.0)
```

<!-- grade_id: cell-d325061167621b36, points: 20, grade: true, solution: false -->
```python
# TEST CASE
true_mean = adult['Capital Gain'].mean()
trials = [saa_avg_capgain(500, 1.0) for _ in range(20)]
errors = [pct_error(true_mean, t) for t in trials]
assert np.mean(errors) < 10
```

## Question 7 (10 points)

In 5-6 sentences, answer the following:

- What clipping values did you choose for clipping the query outputs on each chunk? How do these compare with the best values for clipping individual capital gains values when using the "sum / count" approach?
- Is 500 a good value for the number of chunks $k$? How does making $k$ larger or smaller change the results?
- How does the sample-and-aggregate approach compare to the "sum / count" approach for this particular query?

<!-- grade_id: cell-4bbbc2dba075a1e5, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE