# CS211: Data Privacy
## Homework 3

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')
from unittest.mock import patch

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
adult = adult.dropna()
```

```python
adult.head()
```

## Question 1 (10 points)

Write code to answer the query: "how many participants have never been married?"

*Hint*: filter the `adult_data` dataframe to contain only participants who were never married, then return the  `len` of the filtered dataframe.

<!-- grade_id: cell-975eb03979d78eaf, points: 0, grade: false, solution: true -->
```python
def query1():
    # YOUR CODE HERE
    raise NotImplementedError()

query1()
```

<!-- grade_id: cell-c5ba2c93a46e8c5e, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 1
assert query1() == 9726
```

## Question 2 (5 points)

In 2-5 sentences, answer the following:
- What is the sensitivity of `query1`, and why?

<!-- grade_id: cell-fb1c7b0533f933f7, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 (10 points)

Use the implementation of `laplace_mech` to produce a differentially private answer to `query1`, with `epsilon = 0.1`.

<!-- grade_id: cell-80d3c108ba0f75d5, points: 0, grade: false, solution: true -->
```python
def dp_query1(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

dp_query1(0.1)
```

<!-- grade_id: cell-93eab43d27806309, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 3
dp_results = [dp_query1(0.1) for _ in range(100)]
spec = [np.random.laplace(loc=9726, scale=1/0.1) for _ in range(100)]
assert stats.wasserstein_distance(dp_results, spec) < 10
```

## Question 4 (10 points)

The `pct_error` function, defined below, returns the percent relative error between an original query result and a differentially private result for the same query.

Implement a function `graph_error1` that:

- Calculates 1000 differentially private answers to `dp_query1`
- Calculates the percent error for each one of these answers against the original (non-private) answer
- Graphs the distribution of errors using a histogram

*Hint*: use `plt.hist(..., bins=20)`.

The given code will use your function to plot errors for `epsilon=0.1` and `epsilon=1.0`.

```python
def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0
```

<!-- grade_id: cell-da0c198a1cf9a866, points: 0, grade: false, solution: true -->
```python
def graph_error1(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

graph_error1(0.1)
graph_error1(1.0)
```

<!-- grade_id: cell-eda6bc27840a9067, points: 10, grade: true, solution: false -->
```python
error_spec = 100.0 * np.abs(np.random.laplace(loc=0, scale=1, size=1000)) / query1()

with patch('matplotlib.pyplot.hist') as mock_hist:
    graph_error1(1.0)
    
args, kwargs = mock_hist.call_args
assert stats.wasserstein_distance(error_spec, args[0]) < 0.001
```

## Question 5 (10 points)

In 2-5 sentences, answer the following:

- How does the histogram of relative errors for $\epsilon = 0.1$ differ from the one for $\epsilon = 1.0$?
- What do the two histograms tell you about the effect of $\epsilon$ on relative error?

<!-- grade_id: cell-075013c43cc4cc1d, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 6 (10 points)

Consider `query2`, which asks how many people in the dataset are over the age of 60.

```python
def query2():
    return len(adult[adult['Age'] > 60])
```

Implement `dp_query2`, a differentially private version of `query2` (as in question 3), and `graph_error2`, which graphs relative error for `dp_query2` (as in question 4).

<!-- grade_id: cell-1c15ba7324d807a4, points: 0, grade: false, solution: true -->
```python
def dp_query2(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

def graph_error2(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

graph_error2(1.0)
graph_error1(1.0) # we plot both errors for query 1 and query 2 at the same epsilon, to compare
```

<!-- grade_id: cell-239ae948aa08f924, points: 10, grade: true, solution: false -->
```python
error_spec = 100.0 * np.abs(np.random.laplace(loc=0, scale=1, size=1000)) / query2()

with patch('matplotlib.pyplot.hist') as mock_hist:
    graph_error2(1.0)
    
args, kwargs = mock_hist.call_args
assert stats.wasserstein_distance(error_spec, args[0]) < 0.01
```

## Question 7 (10 points)

In 2-5 sentences, answer the following:

- Given the graph from question 6, how does relative error differ between `dp_query1` and `dp_query2` for the same value of $\epsilon$?
- What property of the query causes the difference in relative errors between `dp_query1` and `dp_query2`?

<!-- grade_id: cell-be5745d92a22154c, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE