# CS211: Data Privacy
## In-Class Exercise, week of 11/07/2022

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def gaussian_mech_vec(vec, sensitivity, epsilon, delta):
    return [v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)
            for v in vec]

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Implement an *encoding function* for randomized response to a "yes/no" question.

Reference [Chapter 13](https://uvm-plaid.github.io/programming-dp/notebooks/ch13.html#randomized-response).

<!-- grade_id: cell-c52b65e148358a20, points: 0, grade: false, solution: true -->
```python
def encode_rand_resp_yes_no(true_response):
    # YOUR CODE HERE
    raise NotImplementedError()

for _ in range(5):
    print('Randomized response:', encode_rand_resp_yes_no(True))
```

<!-- grade_id: cell-5fa94c04030de1ba, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert np.sum([encode_rand_resp_yes_no(True) for _ in range(100)]) > 60
assert np.sum([encode_rand_resp_yes_no(True) for _ in range(100)]) < 90
```

## Question 2

Implement a *decoding function* for randomized response to a "yes/no" question.

<!-- grade_id: cell-3d62f78f67a975f0, points: 0, grade: false, solution: true -->
```python
def decode_rand_resp_yes_no(responses):
    # YOUR CODE HERE
    raise NotImplementedError()

# Example: 1000 "yesses" and 500 "nos"
true_responses = [True for _ in range(1000)] + [False for _ in range(500)]
print('Number of "True" yesses:', np.sum(true_responses))

# Randomized responses
# Each response satisfies 1.09-differential privacy
rand_responses = [encode_rand_resp_yes_no(r) for r in true_responses]

# Decode the responses by subtracting "fake" yesses
print('Decoded randomized response yesses:', decode_rand_resp_yes_no(rand_responses))
```

<!-- grade_id: cell-cc5191292192a8f0, points: 5, grade: true, solution: false -->
```python
# TEST CASE
true_responses = [True for _ in range(1000)] + [False for _ in range(500)]

# Randomized responses
# Each response satisfies 1.09-differential privacy
rand_responses = [encode_rand_resp_yes_no(r) for r in true_responses]

# Decode the responses by subtracting "fake" yesses
assert decode_rand_resp_yes_no(rand_responses) < 1100
assert decode_rand_resp_yes_no(rand_responses) > 900
```

## Question 3

Use the definition of randomized response above to answer the question:

*How many individuals in the `adult` dataset have `Occupation` = `Sales`?*

<!-- grade_id: cell-70d43cb566527293, points: 5, grade: true, solution: true -->
```python
# take a single occupation from the adult dataset, and return a single response
def encode_response_sales(response):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def decode_responses_sales(responses):
    # YOUR CODE HERE
    raise NotImplementedError()

responses = [encode_response_sales(r) for r in adult['Occupation']]
decode_responses_sales(responses)
```

```python
# How accurate is the answer above?
true_sales = np.sum(adult['Occupation'] == 'Sales')
print('True number of salespeople:', true_sales)
```

## Question 4

Implement the *encode* and *perturb* steps for Optimized Unary Hashing.

<!-- grade_id: cell-1809db88535aed7e, points: 0, grade: true, solution: true -->
```python
domain = adult['Occupation'].dropna().unique()
domain

def encode(response):
    # YOUR CODE HERE
    raise NotImplementedError()

encode('Sales')

def perturb(encoded_response):
    # YOUR CODE HERE
    raise NotImplementedError()

perturb(encode('Sales'))
```

```python
# As in randomized response, flipping of bits causes issues in the answers
# This is the perturbed answer (without decoding)
counts = np.sum([perturb(encode(r)) for r in adult['Occupation']], axis=0)
list(zip(domain, counts))
```

```python
# This is the real answer
counts = np.sum([encode(r) for r in adult['Occupation']], axis=0)
list(zip(domain, counts))
```

## Question 5

Implement the *aggregate* step for Optimized Unary Hashing.

<!-- grade_id: cell-0c8cd0a758266161, points: 5, grade: true, solution: true -->
```python
def aggregate(responses):
    # YOUR CODE HERE
    raise NotImplementedError()

responses = [perturb(encode(r)) for r in adult['Occupation']]
counts = aggregate(responses)
list(zip(domain, counts))
```