# CS211: Data Privacy
## In-Class Exercise, 11/02/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Above Threshold

The following code implements Above Threshold, but on a list of *query results* instead of the queries themselves:

```python
# preserves epsilon-differential privacy
def above_threshold(query_results, T, epsilon):
    T_hat = T + np.random.laplace(loc=0, scale = 2/epsilon)
    
    for idx, q in enumerate(query_results):
        nu_i = np.random.laplace(loc=0, scale = 4/epsilon)
        if q + nu_i >= T_hat:
            return idx
    return -1 # the index of the last element
```

## Question 1

Implement `above_threshold_val`, which returns the *value* of the first query result above the threshold. Your solution should have a **total privacy cost** of `epsilon`.

<!-- grade_id: cell-ebb0d6ecbbe9b894, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f7bb22d7edd54872, points: 5, grade: true, solution: false -->
```python
# TEST CASE

results = [above_threshold_val(query_results, 800, 1.0) for _ in range(20)]
assert np.mean(results) > 865
assert np.mean(results) < 890
```

## Question 2

Argue informally that your solution in question 1 has a total privacy cost of `epsilon`.

<!-- grade_id: cell-5e2ea1cc5b8e4c8f, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE