# CS211: Data Privacy
## In-Class Exercise, week of 8/29/2022

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Construct a de-identified version of the `adult` dataset.

<!-- grade_id: cell-4e7a3cdbee680759, points: 0, grade: false, solution: true -->
```python
def deidentify_adult():
    return adult.drop(['Name', 'DOB', 'SSN'], axis=1)

adult_deid = deidentify_adult()
```

```python
deidentify_adult()
```

Output:
```
Zip  Age         Workclass  fnlwgt   Education  Education-Num  \
0      64152   39         State-gov   77516   Bachelors             13   
1      61523   50  Self-emp-not-inc   83311   Bachelors             13   
2      95668   38           Private  215646     HS-grad              9   
3      25503   53           Private  234721        11th              7   
4      75387   28           Private  338409   Bachelors             13   
...      ...  ...               ...     ...         ...            ...   
32556  41328   27           Private  257302  Assoc-acdm             12   
32557  94735   40           Private  154374     HS-grad              9   
32558  49628   58           Private  151910     HS-grad              9   
32559   8213   22           Private  201490     HS-grad              9   
32560  86153   52      Self-emp-inc  287927     HS-grad              9   

           Marital Status         Occupation   Relationship   Race     Sex  \
0           Never-married       Adm-cleric
```

<!-- grade_id: cell-68632b6467ef7923, points: 5, grade: true, solution: false -->
```python
assert 'Name' not in adult_deid.columns
assert 'DOB' not in adult_deid.columns
assert 'SSN' not in adult_deid.columns
```

## Question 2

What columns did you remove from the dataset to de-identify it? Why did you choose these columns? Why did you not choose the remaining columns?

<!-- grade_id: cell-d81a131204234a9d, points: 2, grade: true, solution: true -->
I looked at the test case and saw that I needed to remove Name, DOB, SSN.

These are PII according to HIPAA.

These properties are generally unique or close to unique for individuals

## Question 3

Write a pandas expression to return just the row containing information about Brenn McNeely.

<!-- grade_id: cell-4baa91fd73c5802e, points: 0, grade: false, solution: true -->
```python
def get_brenns_row():
    return adult[adult['Name'] == 'Brenn McNeely']

brenns_row = get_brenns_row()[['Name', 'DOB', 'Zip', 'Age']]
```

<!-- grade_id: cell-13b75b0658ee9750, points: 2, grade: true, solution: false -->
```python
assert len(brenns_row) == 1
assert brenns_row['Zip'].iloc[0] == 95668
assert brenns_row['Age'].iloc[0] == 38
```

```python
brenns_row
```

Output:
```
Name       DOB    Zip  Age
2  Brenn McNeely  8/6/1991  95668   38
```

## Question 4

Conduct a linking attack to recover Brenn's data from the `adult_deid` dataset.

<!-- grade_id: cell-99305ea08eafa471, points: 0, grade: false, solution: true -->
```python
def recover_brenns_row():
    brenns_zip = int(brenns_row['Zip'])
    brenns_age = int(brenns_row['Age'])
    return adult_deid[(adult_deid['Zip'] == brenns_zip) & (adult_deid['Age'] == brenns_age)]

brenns_recovered_row = recover_brenns_row()
brenns_recovered_row
```

Output:
```
Zip  Age Workclass  fnlwgt Education  Education-Num Marital Status  \
2  95668   38   Private  215646   HS-grad              9       Divorced   

          Occupation   Relationship   Race   Sex  Capital Gain  Capital Loss  \
2  Handlers-cleaners  Not-in-family  White  Male             0             0   

   Hours per week        Country Target  
2              40  United-States  <=50K
```

<!-- grade_id: cell-b4432ec45f4faf9c, points: 5, grade: true, solution: false -->
```python
assert len(brenns_recovered_row) == 1
assert brenns_recovered_row['Education'].iloc[0] == 'HS-grad'
```

```python
brenns_recovered_row
```

Output:
```
Zip  Age Workclass  fnlwgt Education  Education-Num Marital Status  \
2  95668   38   Private  215646   HS-grad              9       Divorced   

          Occupation   Relationship   Race   Sex  Capital Gain  Capital Loss  \
2  Handlers-cleaners  Not-in-family  White  Male             0             0   

   Hours per week        Country Target  
2              40  United-States  <=50K
```

```python
adult['Age'].max()
```

Output:
```
90
```

```python
adult['Age'].sum() - adult[adult['Name'] != 'Brenn McNeely']['Age'].sum()
```

Output:
```
38
```

```python
adult['Age'].mean()
```

Output:
```
38.58164675532078
```

```python

```