# CS211: Data Privacy
## In-Class Exercise, 10/26/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

The code below defines a finite set of options for marital status. Define a *scoring function* that returns high scores for common marital statuses, and low scores for uncommon ones (e.g. the score could be the number of people with that status).

<!-- grade_id: cell-d1a5aa9094ce97d0, points: 0, grade: false, solution: true -->
```python
options = ['Never-married', 'Married-civ-spouse', 'Divorced',
           'Married-spouse-absent', 'Separated', 'Married-AF-spouse',
           'Widowed']

def score(option):
    # YOUR CODE HERE
    raise NotImplementedError()

score('Never-married')
```

<!-- grade_id: cell-c4f7745dc6a4cdd2, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert score('Never-married') == 10683
```

## Question 2

Implement `report_noisy_max` using the Laplace mechanism. `report_noisy_max` should return the value in a set that approximately maximizes the value of the score function. It should *not* return the score itself.

<!-- grade_id: cell-8b6fef56dcfaa9ad, points: 0, grade: false, solution: true -->
```python
def report_noisy_max(R, score, sensitivity, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

report_noisy_max(options, score, 1, 1)
```

<!-- grade_id: cell-795dae7b161f7ed7, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert report_noisy_max(options, score, 1, 1) == 'Married-civ-spouse'
```

## Question 3

What is the **total privacy cost** of `report_noisy_max` under *sequential composition*?

<!-- grade_id: cell-20c5b0a19f18423c, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE