# CS211: Data Privacy
## In-Class Exercise, 10/24/2022

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

The code below defines a finite set of options for marital status. Define a *scoring function* that returns high scores for common marital statuses, and low scores for uncommon ones (e.g. the score could be the number of people with that status).

<!-- grade_id: cell-d1a5aa9094ce97d0, points: 0, grade: false, solution: true -->
```python
options = ['Never-married', 'Married-civ-spouse', 'Divorced',
           'Married-spouse-absent', 'Separated', 'Married-AF-spouse',
           'Widowed']

def score(option):
    # YOUR CODE HERE
    raise NotImplementedError()

score('Never-married')
```

<!-- grade_id: cell-c4f7745dc6a4cdd2, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert score('Never-married') == 10683
```

## Question 2

Implement `report_noisy_max` using the Laplace mechanism. `report_noisy_max` should return the value in a set that approximately maximizes the value of the score function. It should *not* return the score itself.

<!-- grade_id: cell-8b6fef56dcfaa9ad, points: 0, grade: false, solution: true -->
```python
def report_noisy_max(R, score, sensitivity, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

report_noisy_max(options, score, 1, 1)
```

<!-- grade_id: cell-795dae7b161f7ed7, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert report_noisy_max(options, score, 1, 1) == 'Married-civ-spouse'
```

## Question 3

What is the **total privacy cost** of `report_noisy_max` under *sequential composition*?

<!-- grade_id: cell-20c5b0a19f18423c, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Above Threshold

The following code implements Above Threshold:

```python
# preserves epsilon-differential privacy
def above_threshold(queries, df, T, epsilon):
    T_hat = T + np.random.laplace(loc=0, scale = 2/epsilon)
    
    for idx, q in enumerate(queries):
        nu_i = np.random.laplace(loc=0, scale = 4/epsilon)
        if q(df) + nu_i >= T_hat:
            return idx
    return -1 # the index of the last element
```

## Question 4

Use `above_threshold` to find the first `age` for which `len(adult[adult['Age'] == age]) >= 800` (the first age for which more than 800 people have that age).

<!-- grade_id: cell-cca35e2da948a4c4, points: 0, grade: false, solution: true -->
```python
def make_query(age):
    def q(df):
        return len(df[df['Age'] == age])
    return q

def find_first_age_above_800(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
    
find_first_age_above_800(1.0)
```

<!-- grade_id: cell-cee279ee8d01cbd0, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert find_first_age_above_800(1.0) == 22
```

## Question 5

Use `above_threshold` to implement `pick_b` for the following summation query. `pick_b` should pick a clipping parameter `b`.

<!-- grade_id: cell-6f208369f5060173, points: 0, grade: false, solution: true -->
```python
def make_query(b):
    def q(df):
        # YOUR CODE HERE
        raise NotImplementedError()
    return q

def pick_b(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
    
pick_b(1.0)
```

<!-- grade_id: cell-9d138a4533e63a07, points: 5, grade: true, solution: false -->
```python
# TEST CASE
b = pick_b(1.0)
assert b > 80
assert b < 100
```

## Question 6

Implement `above_threshold_val`, which returns the *value* of the first query result above the threshold. Your solution should have a **total privacy cost** of `epsilon`.

<!-- grade_id: cell-ebb0d6ecbbe9b894, points: 0, grade: false, solution: true -->
```python
def make_query(age):
    def q(df):
        return len(df[df['Age'] == age])
    return q

def above_threshold_val(queries, df, T, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
    
queries = [make_query(age) for age in range(0,100)]
above_threshold_val(queries, adult, 800, 1.0)
```

<!-- grade_id: cell-f7bb22d7edd54872, points: 5, grade: true, solution: false -->
```python
# TEST CASE
queries = [make_query(age) for age in range(0,100)]

results = [above_threshold_val(queries, adult, 800, 1.0) for _ in range(20)]
assert np.mean(results) > 865
assert np.mean(results) < 890
```

## Question 7

Argue informally that your solution in question 6 has a total privacy cost of `epsilon`.

<!-- grade_id: cell-5e2ea1cc5b8e4c8f, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE