# CS211: Data Privacy
## In-Class Exercise, 10/14/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Write a function to calculate the L1 **local sensitivity** of the following mean query.

```python
def mean_query():
    b = 125
    return adult['Age'][adult['Education'] == 'Bachelors'].clip(lower=0, upper=b).mean()

mean_query()
```

<!-- grade_id: cell-f39828c2d85af776, points: 0, grade: false, solution: true -->
```python
def mean_query_local_sensitivity():
    # YOUR CODE HERE
    raise NotImplementedError()

mean_query_local_sensitivity()
```

<!-- grade_id: cell-f1bb755347650609, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert (mean_query_local_sensitivity() - 0.02333831217326363) < 0.0001
```

## Question 2

Why can't we use the local sensitivity from question 1 directly to achieve differential privacy for `mean_query`?

<!-- grade_id: cell-55d134a1eb26aec9, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE