# CS211: Data Privacy
## In-Class Exercise, 10/11/2021

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Write a function to calculate the L1 **local sensitivity** of the following mean query.

```python
def mean_query():
    b = 125
    return adult['Age'][adult['Education'] == 'Bachelors'].clip(lower=0, upper=b).mean()

mean_query()
```

<!-- grade_id: cell-f39828c2d85af776, points: 0, grade: false, solution: true -->
```python
def mean_query_local_sensitivity():
    # YOUR CODE HERE
    raise NotImplementedError()

mean_query_local_sensitivity()
```

<!-- grade_id: cell-f1bb755347650609, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert (mean_query_local_sensitivity() - 0.02333831217326363) < 0.0001
```

## Question 2

Why can't we use the local sensitivity from question 1 directly to achieve differential privacy for `mean_query`?

<!-- grade_id: cell-55d134a1eb26aec9, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

Implement an upper bound on $A(f, x, k)$, the local sensitivity of a mean query at distance $k$.

<!-- grade_id: cell-bfe45a4f5261718f, points: 0, grade: false, solution: true -->
```python
def ls_mean_at_distance(df, b, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e19c35e52e3a15c0, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert ls_mean_at_distance(adult['Age'], 125, 500) == 0.0038988178784192633
assert ls_mean_at_distance(adult['Age'], 125, 5000) == 0.004535394216465295
```

## Question 4

Implement an upper bound on $D(f, x, s)$, the distance from $x$ to a database where the local sensitivity of the mean query is greater than or equal to $s$.

<!-- grade_id: cell-2750bb4f33a0b420, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-77a206124ddb3baf, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert dist_to_high_ls_mean(adult['Age'], 125, 0.1) == 31311 
assert dist_to_high_ls_mean(adult['Age'], 125, 0.01) == 20061
assert dist_to_high_ls_mean(adult['Age'], 125, 0.001) == 0
```

## Question 5

Use the Propose-Test-Release (PTR) framework to calculate the mean age in the adult dataset with differential privacy.

<!-- grade_id: cell-62e95c3348999e9a, points: 0, grade: false, solution: true -->
```python
def mean_ptr(epsilon, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

mean_ptr(1.0, 10e-5)
```

<!-- grade_id: cell-efb8d6da908b8ae9, points: 5, grade: true, solution: false -->
```python
# TEST CASE

true_answer = adult['Age'].mean()
trials = [mean_ptr(1.0, 10e-5) for _ in range(20)]
errors = [pct_error(true_answer, x) for x in trials]
assert np.mean(errors) < 1.0
```

## Exploration: how low can the proposal be?

```python
xs = np.linspace(0.001, 0.01)
clipped_ages = adult['Age'].clip(lower=0, upper=125)
ys = [dist_to_high_ls_mean(clipped_ages, 125, p) for p in xs]
```

```python
plt.plot(xs, ys)
```