# CS211: Data Privacy
## In-Class Exercise, 9/23/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Write a function `dp_mean_age` that calculates the mean age of members of the dataset.

```python
def mean_age():
    return adult['Age'].mean()

mean_age()
```

<!-- grade_id: cell-8dd5afd2295449ac, points: 0, grade: false, solution: true -->
```python
def dp_mean_age(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-460c4e11928bdfa5, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1
results = [dp_mean_age(1) for _ in range(100)]
assert np.abs(np.mean(results) - mean_age()) < 0.1
```

## Question 2

Write a function `dp_age_sum_clipping` that uses clipping to bound the sensitivity of the sum of ages.

<!-- grade_id: cell-0ad1c2c480c24c99, points: 0, grade: false, solution: true -->
```python
def dp_age_sum_clipping(epsilon, b):
    # YOUR CODE HERE
    raise NotImplementedError()

dp_age_sum_clipping(1.0, 100)
```

<!-- grade_id: cell-6f0418293b0ed25c, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 2
results = [dp_age_sum_clipping(1.0, 100) for _ in range(10)]
assert np.abs(np.mean(results) - adult['Age'].sum()) < 100
```

```python
bs = np.linspace(1, 125, 100)
results = [dp_age_sum_clipping(0.1, b) for b in bs]
plt.plot(bs, results);

#baseline = [adult['Age'].sum() for b in bs]
#plt.plot(bs, baseline);
```