# CS211: Data Privacy
## In-Class Exercise, 9/11/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Write code to check if a dataframe satisfies $k$-Anonymity for a given $k$.

<!-- grade_id: cell-c2f5da8fc63eb982, points: 0, grade: false, solution: true -->
```python
def isKAnonymized(df, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-dfebf0863cfc1459, points: 5, grade: true, solution: false -->
```python
# Test cases for question 1

df1 = adult[['Education']].dropna()
assert isKAnonymized(df1[:1000], 3) == False
assert isKAnonymized(df1[:1000], 2) == True
assert isKAnonymized(df1[:1000], 1) == True

df2 = adult[['Zip', 'Education']].dropna()
assert isKAnonymized(df2[:1000], 3) == False
assert isKAnonymized(df2[:1000], 2) == False
assert isKAnonymized(df2[:1000], 1) == True
```

## Question 2

What is the time complexity of the algorithm in Question 1?

<!-- grade_id: cell-2b613e79c39137ae, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

Write code to generalize the `Zip` column in the adult dataset.

<!-- grade_id: cell-be9a4264a0e58938, points: 0, grade: false, solution: true -->
```python
def generalize_zip(zip, l):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-256688fe78b8f7bf, points: 5, grade: true, solution: false -->
```python
# Test cases for question 3
assert generalize_zip(47401, 0) == 47401
assert generalize_zip(47401, 1) == 47400
assert generalize_zip(47401, 2) == 47400
assert generalize_zip(47401, 3) == 47000
assert generalize_zip(47401, 4) == 40000
assert generalize_zip(47401, 5) == 0
```