# CS211: Data Privacy
## In-Class Exercise, 9/2/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1 (5 points)

Write code to group the `adult` dataset by a single column and count the number of members in each group.

<!-- grade_id: cell-055399c92e676783, points: 0, grade: false, solution: true -->
```python
def group_one_count(col):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
group_one_count('Education')
```

<!-- grade_id: cell-befc11c8be0bf91b, points: 5, grade: true, solution: false -->
```python
s = group_one_count('Education')
assert s['10th'] == 933
assert s['9th'] == 514
assert s['Some-college'] == 7291
```

## Question 2 (5 points)

Write code to group the `adult` dataset by two columns and count the number of members in each group.

<!-- grade_id: cell-f6a37efe405c4a98, points: 0, grade: false, solution: true -->
```python
def group_two_count(col1, col2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
group_two_count('Occupation', 'Education')
```

<!-- grade_id: cell-88caedc5d6e44568, points: 5, grade: true, solution: false -->
```python
s = group_two_count('Occupation', 'Education')
assert s['Transport-moving']['Doctorate'] == 1
assert s['Adm-clerical']['10th'] == 38
```

## Question 3 (5 points)

Write code to perform a differencing attack to determine Brenn McNeely's age using only the `mean` aggregation function over large groups.

<!-- grade_id: cell-6fb63040052b0857, points: 0, grade: false, solution: true -->
```python
def get_brenns_age():
    # YOUR CODE HERE
    raise NotImplementedError()

brenns_age = get_brenns_age()
```

<!-- grade_id: cell-d66b7d6c32d256c6, points: 5, grade: true, solution: false -->
```python
assert brenns_age == 38.0
```