# CS211: Data Privacy
## In-Class Exercise, 9/18/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Implement the Laplace mechanism below.

<!-- grade_id: cell-9f17a2616172b172, points: 0, grade: false, solution: true -->
```python
def laplace_mech(v, sensitivity, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-cae8dc00eebe8255, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1
dist1 = [laplace_mech(50, 1, 1.0) for _ in range(200)]
dist2 = [np.random.laplace(loc=50, scale=1) for _ in range(200)]

assert stats.wasserstein_distance(dist1, dist2) < 1
```

## Question 2



<!-- grade_id: cell-9439553739156776, points: 0, grade: false, solution: true -->
```python
def people_over_30_dp():
    # YOUR CODE HERE
    raise NotImplementedError()

people_over_30_dp()
```

<!-- grade_id: cell-6b8b93375f070186, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 2
dist1 = [np.random.laplace(loc=21989, scale=1.0) for _ in range(200)]
dist2 = [people_over_30_dp() for _ in range(200)]

assert stats.wasserstein_distance(dist1, dist2) < 1
```