# CS211: Data Privacy
## In-Class Exercise, 10/19/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Use the Propose-Test-Release (PTR) framework to calculate the mean age in the adult dataset with differential privacy.

<!-- grade_id: cell-62e95c3348999e9a, points: 0, grade: false, solution: true -->
```python
def ls_mean_at_distance(df, b, k):
    return np.abs(b/(len(df) - k))

def dist_to_high_ls_mean(df, b, s):
    k = 0
    
    while ls_mean_at_distance(df, b, k) < s:
        k += 1
    
    return k
    
def mean_ptr(epsilon, delta):
    # YOUR CODE HERE
    raise NotImplementedError()

mean_ptr(1.0, 10e-5)
```

<!-- grade_id: cell-efb8d6da908b8ae9, points: 5, grade: true, solution: false -->
```python
# TEST CASE

true_answer = adult['Age'].mean()
trials = [mean_ptr(1.0, 10e-5) for _ in range(20)]
errors = [pct_error(true_answer, x) for x in trials]
assert np.mean(errors) < 1.0
```

## Exploration: how low can the proposal be?

```python
xs = np.linspace(0.001, 0.01)
clipped_ages = adult['Age'].clip(lower=0, upper=125)
ys = [dist_to_high_ls_mean(clipped_ages, 125, p) for p in xs]
plt.plot(xs, ys);
```

## Question 2

Use the sample and aggregate framework to release the average age in the `adult` dataset with differential privacy. [See reference here](https://uvm-plaid.github.io/programming-dp/notebooks/ch7.html#sample-and-aggregate).

<!-- grade_id: cell-1844b42041061beb, points: 0, grade: false, solution: true -->
```python
def f(series):
    return series.mean()

def saa_avg_age(k, epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

saa_avg_age(200, 1)
```

<!-- grade_id: cell-9a6b4d98ce3b93f9, points: 5, grade: true, solution: false -->
```python
# TEST CASE

true_answer = adult['Age'].mean()
trials = [saa_avg_age(200, 1.0) for _ in range(20)]
errors = [pct_error(true_answer, x) for x in trials]
assert np.mean(errors) < 3.0
```