# CS211: Data Privacy
## In-Class Exercise, 10/30/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def gaussian_mech(v, sensitivity, epsilon, delta):
    return v + np.random.normal(loc=0, scale=sensitivity * np.sqrt(2*np.log(1.25/delta)) / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Above Threshold

The following code implements Above Threshold:

```python
# preserves epsilon-differential privacy
def above_threshold(queries, df, T, epsilon):
    T_hat = T + np.random.laplace(loc=0, scale = 2/epsilon)
    
    for idx, q in enumerate(queries):
        nu_i = np.random.laplace(loc=0, scale = 4/epsilon)
        if q(df) + nu_i >= T_hat:
            return idx
    return -1 # the index of the last element
```

## Question 1

Use `above_threshold` to find the first `age` for which `len(adult[adult['Age'] == age]) >= 800` (the first age for which more than 800 people have that age).

<!-- grade_id: cell-cca35e2da948a4c4, points: 0, grade: false, solution: true -->
```python
def make_query(age):
    def q(df):
        return len(df[df['Age'] == age])
    return q

def find_first_age_above_800(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
    
find_first_age_above_800(1.0)
```

<!-- grade_id: cell-cee279ee8d01cbd0, points: 5, grade: true, solution: false -->
```python
# TEST CASE
assert find_first_age_above_800(1.0) == 22
```

## Question 2

Use `above_threshold` to implement `pick_b` for the following summation query. `pick_b` should pick a clipping parameter `b`.

<!-- grade_id: cell-6f208369f5060173, points: 0, grade: false, solution: true -->
```python
def make_query(b):
    def q(df):
        # YOUR CODE HERE
        raise NotImplementedError()
    return q

def pick_b(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()
    
pick_b(1.0)
```

<!-- grade_id: cell-9d138a4533e63a07, points: 5, grade: true, solution: false -->
```python
# TEST CASE
b = pick_b(1.0)
assert b > 80
assert b < 100
```