# CS211: Data Privacy
## In-Class Exercise, 9/18/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Consider the following definition of a differencing attack (without differential privacy).

```python
def differencing_attack():
    q1 = adult['Age'].sum()
    q2 = adult[adult['Name'] != 'Karrie Trusslove']['Age'].sum()

    return q1 - q2

print('Differencing attack result:', differencing_attack())
```

Implement a function `dp_differencing_attack` that performs the same attack, but attempts to satisfy differential privacy using the Laplace mechanism. Use the parameters `sensitivity=1` and `epsilon=1.0`.

<!-- grade_id: cell-df30d40b4e25882d, points: 0, grade: false, solution: true -->
```python
def dp_differencing_attack():
    q1 = adult['Age'].sum()
    q2 = adult[adult['Name'] != 'Karrie Trusslove']['Age'].sum()
    
    # YOUR CODE HERE
    raise NotImplementedError()

print('DP Differencing attack result:', dp_differencing_attack())
```

<!-- grade_id: cell-cbc91dc6c256bc87, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1
dp_results = [dp_differencing_attack() for _ in range(100)]
spec = [np.random.laplace(loc=39, scale=1) for _ in range(100)]
assert stats.wasserstein_distance(dp_results, spec) < 2
```

## Question 2

What is the *sensitivity* of the `differencing_attack` query defined above, and why?

<!-- grade_id: cell-f7a4f9087111dbfd, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

Implement a corrected version of `dp_differencing_attack` that uses the correct sensitivity and thus correctly satisfies differential privacy.

<!-- grade_id: cell-ddf7057ec6927a23, points: 0, grade: false, solution: true -->
```python
def dp_differencing_attack_corrected():
    q1 = adult['Age'].sum()
    q2 = adult[adult['Name'] != 'Karrie Trusslove']['Age'].sum()
    
    # YOUR CODE HERE
    raise NotImplementedError()

print('DP Differencing attack result:', dp_differencing_attack_corrected())
```

<!-- grade_id: cell-626edd0f2a305401, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 3
dp_results = [dp_differencing_attack_corrected() for _ in range(100)]
spec = [np.random.laplace(loc=39, scale=125/2.0) for _ in range(100)]
assert stats.wasserstein_distance(dp_results, spec) < 200
assert stats.wasserstein_distance(dp_results, spec) > 50
```