# CS211: Data Privacy
## In-Class Exercise, 9/25/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

For various values of $b$, write code to print out the percent error of summing the ages in the `adult` dataset, 
clipped to each value of $b$.

<!-- grade_id: cell-bb6f4fe4b3beb456, points: 0, grade: false, solution: true -->
```python
bs = range(1, 100, 10)
real_sum = adult['Age'].sum()

# YOUR CODE HERE
raise NotImplementedError()
```

What value of $b$ is the best?

<!-- grade_id: cell-250b062d06204617, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2

For various values of $b$, print the result of a *differentially private* sum of ages, clipped to each value of $b$. Use $\epsilon = 0.1$.

<!-- grade_id: cell-dca4f1bfa29d353d, points: 0, grade: false, solution: true -->
```python
bs = range(1, 100, 10)
real_sum = adult['Age'].sum()

# YOUR CODE HERE
raise NotImplementedError()
```

Which value of $b$ is the best now? Does it differ?

<!-- grade_id: cell-fee4808b879b0a04, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

Write an algorithm to *automatically pick the clipping parameter* for the age column. Your algorithm should satisfy differential privacy.

<!-- grade_id: cell-9a41939138d41f6f, points: 0, grade: false, solution: true -->
```python
def pick_b(epsilon):
    # YOUR CODE HERE
    raise NotImplementedError()

pick_b(1.0)
```

<!-- grade_id: cell-136e29134703a26b, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 3

many_trials = np.mean([pick_b(1.0) for _ in range(100)])
assert many_trials > 70
assert many_trials < 500
```

## Question 4

What is the privacy cost of your algorithm, and why?

<!-- grade_id: cell-5e3ac9bbd3f179e9, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE