# CS211: Data Privacy
## In-Class Exercise, 9/14/2020

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Write a *counting query* to determine whether or not Karrie Trusslove's age is 39.

<!-- grade_id: cell-e9c0c06eb7735ea3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f93cfd0eb59130ed, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 1

assert karrie_query() == 1
```

## Question 2

Add Laplace noise to the counting query you wrote in question 1 to ensure differential privacy for $\epsilon = 1.0$.

<!-- grade_id: cell-8ad72c55497bc534, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a7c9bd084e5c5bdf, points: 5, grade: true, solution: false -->
```python
# TEST CASE for question 2

q2_runs = [dp_karrie_query() for _ in range(100)]
noise_runs = [np.random.laplace(loc=1, scale=1/epsilon) for _ in range(100)]

assert stats.wasserstein_distance(q2_runs, noise_runs) < 1
```

## Question 3

In 2-4 sentences, describe how `dp_karrie_query` protects Karrie's privacy.

<!-- grade_id: cell-07337a085c0ba0e1, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE